# Week 5 Exercises

_McKinney 3.2_

**Unlike in previous weeks, in the exercises below, you will need to create a function definition from scratch.  I'll provide specific instructions and examples for how it will be used, but you will have to do the work of definiging the whole function.**

---
---

### 25.1 Longest String

Write a function called `longest (L)` that takes as its only parameter a list of strings.  Your function needs to find the longest of those strings and return the position number of that longest string.  An example is provided below:

The longest word in that list of strings is "birthday" in position #1, so `longest(strings)` should return 1.
```
>>> strings = ['happy', 'birthday', 'to', 'me']
>>> longest(strings)
1
```

Be sure to include a docstring.  Include test cases in your docstring if you want to.

<!-- grade_id: longest-answer, points: 0, grade: false, solution: true -->
```python
def longest(strings_list):
    """
    (list) -> int
    strings_list is a list of strings. This function returns the index of the longest string in the list.

    Examples:
    >>> longest(['happy', 'birthday', 'to', 'me'])
    1
    >>> longest(['enjoy', 'class'])
    0
    >>> longest(['when', 'what', 'where', 'how', 'who'])
    2
    """
    # Create a list of the lengths of each string in strings_list
    lengths = [len(string) for string in strings_list]

    # Find the index of the longest string by locating the index of the max length
    longest_index = lengths.index(max(lengths))

    # Return the index of the longest string
    return longest_index

```

<!-- grade_id: longest-tests, points: 10, grade: true, solution: false -->
```python
assert longest("happy birthday to me".split(" ")) == 1
assert longest("enjoy class".split(" ")) == 0
assert longest(['when','what','where','how','who']) == 2
```

```python
import doctest
doctest.run_docstring_examples(longest, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    longest(['happy', 'birthday', 'to', 'me'])
Expecting:
    1
ok
Trying:
    longest(['enjoy', 'class'])
Expecting:
    0
ok
Trying:
    longest(['when', 'what', 'where', 'how', 'who'])
Expecting:
    2
ok
```

### 25.2 Farenheit to Celsius

Write a function called f_to_c() that converts a given temperature in degrees Farenheit to degrees Celsius.  If you don't recall that conversion, it is:

$ temp_c = \frac{5}{9} \times (temp_f - 32) $

Make sure that you code is well documented using the DocString examples, and that your code includes tests for 212f, 32f, and 98.6f.

```python
def f_to_c(temp_in_f):
    """
    (float) -> float

    temp_in_f is the temperature in degrees Fahrenheit.

    This function converts the given Fahrenheit temperature to Celsius and returns the result.

    Examples:
    >>> f_to_c(212)
    100.0
    >>> f_to_c(32)
    0.0
    >>> f_to_c(98.6)
    37.0
    """
    # Convert Fahrenheit to Celsius using the formula: (tempf - 32) * 5/9
    temp_in_c = (temp_in_f - 32) * 5/9

    # Return the calculated Celsius temperature
    return temp_in_c

```

```python
assert f_to_c(212) == 100.0
assert f_to_c(32) == 0.0
assert f_to_c(98.6) == 37.0
```

```python
import doctest
doctest.run_docstring_examples(f_to_c, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    f_to_c(212)
Expecting:
    100.0
ok
Trying:
    f_to_c(32)
Expecting:
    0.0
ok
Trying:
    f_to_c(98.6)
Expecting:
    37.0
ok
```

### 25.3 Computing Length of Stay

For this problem, we have a collection of patient enounter data stored as a Python dictionary.  The `key` for the dictionary is the **encounter ID**, a code that starts with the letter `E` followed by four numbers.  The value associated with each encounter ID is another Python dictionary.  This "inner" dictionary holds three items: admit date, primary diagnosis, and discharge date.  See the example in the code below.

You need to write a length of stay function that computes the length of stay, in whole days, between the admit date and discharge date.  However, if the diagnosis is "Observation" then the length of stay should always be returned as 0 regardless of the admit and discharge dates. Your function should should take three parameters: a start date, an end date, and a diagnosis.

You will find it handy to refer to this example here on how to calculate the number of days between two dates: https://stackoverflow.com/questions/151199/how-to-calculate-number-of-days-between-two-given-dates

```python
from datetime import date

def los(starting_date, ending_date, diagnosis):
    """
    (date, date, str) -> int
    * starting_date: admission date (date object)
    * ending_date: discharge date (date object)
    * diagnosis: primary diagnosis for the patient (string)

    This function calculates the length of stay (LOS) in full days based on the admission and discharge dates.
    If the diagnosis is 'Observation', the function returns 0, meaning the stay is not counted.

    Examples:
    >>> los(date(2019,1,3), date(2019,1,8), "COPD")
    5
    >>> los(date(2020,1,5), date(2020,1,8), "Observation")
    0
    """
    # Check if the diagnosis is 'Observation'. If so, return 0, indicating no stay.
    if diagnosis == 'Observation':
        return 0

    # Calculate and return the length of stay as the difference between end_date and start_date in days
    return (ending_date - starting_date).days

```

```python
assert los(date(2019,1,3), date(2019,1,8), "COPD") == 5
assert los(date(2020,1,5), date(2020,1,8), "Observation") == 0
```

```python
import doctest
doctest.run_docstring_examples(los, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    los(date(2019,1,3), date(2019,1,8), "COPD")
Expecting:
    5
ok
Trying:
    los(date(2020,1,5), date(2020,1,8), "Observation")
Expecting:
    0
ok
```

```python
encounters = {
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}


for encid, visit in encounters.items():
    print(los(visit['admit'],visit['discharge'],visit['diagnosis']))

# Your output should be
# 5
# 4
# 1
# 0
# 0
```

Output:
```
5
4
1
0
0
```

---
---

### 25.4 Average Length of Stay

Create a function called `average_los` that returns the average of the LOS for the encounters in the provided dictionary. The encounters must be a dictionary that contains dictionaries that each contain at least an admit date and a discharge date.
    
If the length of stay of any individual encounter is 0, then it will not be counted toward the average.

```python
from datetime import date

def calculate_los(admitting_date, discharging_date):
    """
    (date, date) -> int
    Admitting_date: the admission date (date object)
    Discharging_date: the discharge date (date object)

    This function calculates the length of stay (LOS) in whole days based on the provided admitting and discharging dates.

    Examples:
    >>> calculate_los(date(2019,1,3), date(2019,1,8))
    5
    >>> calculate_los(date(2020,1,5), date(2020,1,8))
    3
    """
    # Calculate and return the difference in days between discharging_date and admitting_date
    return (discharging_date - admitting_date).days

def average_los(encounters):
    """
    (dict) -> float
    * encounters: a dictionary where each key is an encounter ID and each value is another dictionary
      with keys 'admit', 'discharge', and 'diagnosis'.

    This function calculates the average length of stay (LOS) for the provided encounters.
    Encounters with a length of stay of 0 or a diagnosis of 'Observation' are excluded from the average.

    Examples:
    >>> encounters = {
    ...     "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    ...     "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    ...     "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) }
    ... }
    >>> round(average_los(encounters), 2)
    4.5
    """
    # Initialize total_los to accumulate total length of stay and count to track valid encounters
    total_los = 0
    count = 0

    # Loop through each encounter in the encounters dictionary
    for encounter_id, data in encounters.items():
        # Extract admitting date, discharging date, and diagnosis for each encounter
        admitting_date = data.get('admit')
        discharging_date = data.get('discharge')
        diagnosis = data.get('diagnosis')

        # Ensure both dates are present and the diagnosis is not 'Observation'
        if admitting_date and discharging_date and diagnosis != 'Observation':
            # Calculate the length of stay using the calculate_los function
            los = calculate_los(admitting_date, discharging_date)

            # Only include encounters with a positive length of stay
            if los > 0:
                total_los += los  # Add the length of stay to the total
                count += 1  # Increment the valid encounter count

    # If no valid encounters, return 0.0 to avoid division by zero
    if count == 0:
        return 0.0

    # Return the average length of stay
    return total_los / count

```

```python
from datetime import date
encounters = {
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}

assert(round(average_los(encounters),2) == 3.33)
```

```python
import doctest
doctest.run_docstring_examples(average_los, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    encounters = {
        "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
        "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
        "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) }
    }
Expecting nothing
ok
Trying:
    round(average_los(encounters), 2)
Expecting:
    4.5
ok
```

---

### 25.5 Celsius to Farenheit

Write your own function called `c_to_f` that converts degrees Celsius to degrees Farenheit.  Include in your solution a series of doc tests that can verify the conversion using inputs of 100, 0, and 37 degrees Celsius

```python
def c_to_f(degrees_celsius):
    """
    (float) -> float
    * degrees_celsius: temperature in degrees Celsius (float)

    This function converts the given Celsius temperature to Fahrenheit and returns the result.

    Examples:
    >>> assert round(c_to_f(100), 1) == 212.0
    >>> assert round(c_to_f(0), 1) == 32.0
    >>> assert round(c_to_f(37), 1) == 98.6
    """
    # Convert Celsius to Fahrenheit using the formula: F = (C * 9/5) + 32
    degrees_fahrenheit = (degrees_celsius * 9/5) + 32

    # Return the calculated temperature in Fahrenheit
    return degrees_fahrenheit

```

```python
import doctest
doctest.run_docstring_examples(c_to_f, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    assert round(c_to_f(100), 1) == 212.0
Expecting nothing
ok
Trying:
    assert round(c_to_f(0), 1) == 32.0
Expecting nothing
ok
Trying:
    assert round(c_to_f(37), 1) == 98.6
Expecting nothing
ok
```

```python

```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

First run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Then open a new command line by clicking the `+` icon above the file list and chosing `Terminal`
3. At the command line in the new Terminal, follow these steps:
  1. Change directories to your project folder and the week03 subfolder (`cd <folder name>`)
  2. Make sure your project folders are up to date with github.com (`git pull`)
  3. Add the homework files for this week (`git add <file name>`)
  4. Commit your changes (`git commit -a -m "message"`)
  5. Push your changes (`git push`)
  
If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.