<a href="https://colab.research.google.com/github/meghanavallam123/HD5210Homework/blob/main/week02/week02_assignment_2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Week 2 Programming Exercises

See: _McKinney 2.3_ and [Python Documentation](https://docs.python.org/3/tutorial/introduction.html) section 3.1.1 on using Python as a calculator.

**At the begining of the semester, all of the workshop programming exercises will be structured a specific way to make it easier to verify for yourself that you're on the right track as well as easier for me to do a first pass on automated grading.  The structure of each question will require you to write a function using Python code. Don't worry that we haven't talked about functions yet. Just edit the code between** `### BEGIN SOLUTION` and `### END SOLUTION` **as shown in the example below.**


**WHAT I PROVIDE:**
```
def some_function(parameter1, parameter2):

   ### BEGIN SOLUTION
   x = -1
   ### END SOLUTION
   
   return x
```

**WHAT YOU SHOULD DO:** Just change the parameter names (if you feel you need to) and the calculations between `### BEGIN SOLUTION` and `### END SOLUTION`.  This is just a made up example.
```
def some_function(a, b):

   ### BEGIN SOLUTION
   temp = a + b
   x = temp / a * b
   ### END SOLUTION
   
   return x
```


---
---

**Below each programming exercise are some tests (`assertions`) that verify your code is working correctly.  If any assertions fail, you know that something isnt' right with your code, but having all assertions pass doesn't necessarily mean your code is perfect, yet. You should also create your own tests to make sure your code is correct.**

**For now, please don't change any function names**

### 5.1 BMI Calculation

Create new names for the variables that you need in the code framework below and use the variables to compute the corresponding body mass index (BMI). You can use any meaningful variable names you like. When you compute the BMI, assign a variable to it, and then print the result to the notebook's output.

* Weght in kilograms is 90.7
* Height in meters in 1.82

Note that a simple BMI calculation is weight in kilograms divided by the square of height in meters.

$ bmi = \frac{weight\,\mathrm{kg}}{(height\,\mathrm{m})^2} $

For this problem, you'll be graded both on proper calculation as well as on programming style and appropriate variables names.

_Also note that BMI should not be considered a universal measure of health or fitness. A health BMI can vary greatly depending on many other fitness conditions and body types._

<!-- grade_id: bmi-answer, points: 0, grade: false, solution: true -->
```python
def calculate_bmi(weight_kg, height_m):

    # 1. Change the names of parameter1 and parameter2 to be something more meaningful
    # 2. Change the name of the return variable below

    ### BEGIN SOLUTION
    bmi_value = weight_kg / (height_m ** 2)
    ### END SOLUTION


    return bmi_value
```

Test your code below.  The BMI for a weight of 90.7 kg and height of 182 cm should be approximately 27.38

<!-- grade_id: bmi-test, points: 0, grade: true, solution: false -->
```python
assert round(calculate_bmi(90.7, 1.82), 2) == 27.38
assert round(calculate_bmi(91, 1.80), 2) == 28.09
assert round(calculate_bmi(70, 1.70), 2) == 24.22
```

```python
bmi_result = calculate_bmi(90.7, 1.82)


print(bmi_result)
```

Output:
```
27.38195870064002
```

### 5.2 Weight from BMI

For the next operation, we're going to calculate Weight in kilograms, given someone's BMI and their height in meters.

This calculation will be:  $ weight\,\mathrm{kg} = {bmi}\times{(height\,\mathrm{m})^2} $

```python
def calculate_weight(bmi, height_m):

    ### BEGIN SOLUTION
    weight_kg = bmi * (height_m ** 2)
    ### END SOLUTION


    return weight_kg
```

```python
assert round(calculate_weight(27.38, 1.82),2) == 90.69
assert round(calculate_weight(28.09, 1.80),2) == 91.01
assert round(calculate_weight(24.22, 1.70),2) == 70.00
```

```python
bmi = 27.38  # BMI value
height = 1.82  # height in meters

weight = calculate_weight(bmi, height)
print(f"The weight for a BMI of {bmi} and height of {height} m is: {weight:.2f} kg")
```

Output:
```
The weight for a BMI of 27.38 and height of 1.82 m is: 90.69 kg
```

### 5.3 Height from BMI


For the next operation, we're going to calculate height in meters, given someone's BMI and their weight in kilograms.

This calculation will be:  $ {height\,\mathrm{m}} = \sqrt{\frac{weight\,\mathrm{kg}}{bmi}} $

For this equation, you'll notice that we need to calculate a square root.  Python has a function for this of course, `sqrt`, but that function is not part of the base Python language. Some Python features have to be imported.  The first line of code below imports the `sqrt` function from the [math](https://docs.python.org/3/library/math.html) library.


```python
from math import sqrt

def height(bmi, weight_kg):

    ### BEGIN SOLUTION
    height_m = sqrt(weight_kg / bmi)
    ### END SOLUTION

    return  height_m
```

```python
assert round(height(27.38, 90.69),2) == 1.82
assert round(height(28.09, 91.01),2) == 1.80
assert round(height(24.22, 70.00),2) == 1.70
```

```python
height_result = height(28.09, 91.01)
print(height_result)
```

Output:
```
1.7999841777715266
```

### 5.4 BMI using Feet, Inches, and Pounds

Now, let's try to to create a function that knows how to first convert between imperial units (US inches and pounds) and SI units before calculating the bmi.  For these calculations use:

* $ 1\,\mathrm{inch} = 2.54\,\mathrm{cm} $
* $ 1\,\mathrm{lb} = 0.45454545\,\mathrm{kg} $


```python
def bmi_us(weight_lb, height_ft, height_in):

    ### BEGIN SOLUTION
    inches_to_cm = 2.54
    lbs_to_kg = 0.45454545


    # Convert height to centimeters and then to meters
    total_height_in_cm = (height_ft * 12 + height_in) * inches_to_cm
    height_m = total_height_in_cm / 100  # Convert centimeters to meters

    # Convert weight to kilograms
    weight_kg = weight_lb * lbs_to_kg

    # Calculate BMI
    bmi_value = weight_kg / (height_m ** 2)
    return bmi_value



```

```python
assert round(bmi_us(250,6,4),2) == 30.49
assert round(bmi_us(120,5,0),2) == 23.48
assert round(bmi_us(130,5,10),2) == 18.69

bmi_result = bmi_us(250, 6, 4)
print(bmi_result)

```

Output:
```
30.494583474450884
```

### 5.5 Target BMI
Assuming that someone doesn't change height over time, what is the approximate number of days it will take for someone to get from the BMI you found in #3 above to a BMI of 24.9 (typically regarded as being in a healthy range) while losing weight at a rate of 0.75 kg per week.

As needed, use variables and multiple calculation steps to make your thought process visible in the code. For instance, rather than using the notebook as a calculator, use variables for intermediate answers.

Again, you will receive points both for correct answers as well as for style.

In addition to the bmi, height, and weight equations we already defined above, here is some of the algebra to help you along:

$ height\,\mathrm{m} = \frac{height\,\mathrm{cm}}{100} $

$ days = {weeks} x {7} $

$ weeks\,required\,to\,change\,weight = \frac{current\,weight - target\,weight}{weight\,change\,per\,week} $

<!-- grade_id: weight-answer, points: 0, grade: false, solution: true -->
```python
from math import sqrt

def lose_weight(weight, height_cm, kg_per_week, target_bmi):
    # Convert height from cm to meters
    height_m = height_cm / 100

    # Calculate the target weight for the desired BMI
    target_weight = target_bmi * (height_m ** 2)

    # Calculate the weight change needed
    weight_change_needed = weight - target_weight

    # Calculate the number of weeks required to reach the target weight
    weeks_required = weight_change_needed / kg_per_week

    # Convert weeks to days
    days = weeks_required * 7

    return days

```

Test your code below.  It should return approximately 76.73 days for the given parameters above.

<!-- grade_id: weight-tests, points: 5, grade: true, solution: false -->
```python
assert round(lose_weight(90.7, 182, 0.75, 24.9),2) == 76.73
assert round(lose_weight(91, 180, 0.75, 24.9),2) == 96.36
assert round(lose_weight(70, 170, 0.75, 22),2) == 59.92
assert round(lose_weight(70, 170, 2, 22),2) == 22.47

```

```python
days_result = lose_weight(90.7, 182, 0.75, 24.9)
print(days_result)
```

Output:
```
76.73157333333342
```

---

### 5.6 Your choice

Every one of us has done projects where some kind of simple mathematical calculators are required.  Think about something recent from your field of expertise or an area of interest. Choose something that you find interesting and write the Python code required to do that calculation.  These examples show the level of complexity you need to have (not very complex).

* $ bmr_{male} = 66 + (13.7 \times weight\,\mathrm{kg}) + (5 \times height\,\mathrm{m}) - (6.8 \times age\,\mathrm{yr}) $

* $ bmr_{female} = 655 + (9.6 \times weight\,\mathrm{kg}) + (1.8 \times height\,\mathrm{cm}) - (4.7 \times age\,\mathrm{yr}) $

* $ Batting\,Average = \frac{Total\,Hits}{Total\,at\,Bats} $

* $ Area\,of\,circle = \pi r^2 $

* $ F = \frac{1}{2} m v^2 $

Then provide the code for how you would use variables and example values to do the calculation.  You do not to wrap this in a Python function. Just include the code and variables as I've shown below with the BMI calculation that we've been using. As the last command in the cell, print out the value of your calculation as shown in my example.


```python
import math  # Import math module to use the value of pi

# Example values for BMR calculation
weight_kg = 70           # Weight in kilograms
height_cm = 165          # Height in centimeters
age_yr = 30              # Age in years

# Calculate BMR for a female using the formula
bmr_female = 655 + (9.6 * weight_kg) + (1.8 * height_cm) - (4.7 * age_yr)

# Print the BMR result
print("BMR for female:", bmr_female)
```

Output:
```
BMR for female: 1483.0
```

```python
# Calculate the area of a circle with a radius of 5 cm
radius_cm = 5
area_circle = math.pi * (radius_cm ** 2)

# Print the area of the circle
print("Area of the circle:", area_circle)
```

Output:
```
Area of the circle: 78.53981633974483
```

```python
# Calculate the force with mass of 2 kg and velocity of 3 m/s
mass_kg = 2
velocity_m_s = 3
force = 0.5 * mass_kg * (velocity_m_s ** 2)

# Print the force
print("Force:", force)
```

Output:
```
Force: 9.0
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to save this notebook file back to GitHub.  To do that in Google Colab:
1. File -> Save a Copy in GitHub
2. Make sure your HDS5210 repository is selected
3. Make sure the file name includes the week number like this: `week02/week02_assignment_2.ipynb`
4. Add a commit message that means something

**Be sure week names are lowercase and use a two digit week number!!**

**Be sure you use the same file name provided by the instructor!!**

